# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In this GitHub repository, there is a file called `allergies.json` that contains a list of patient allergies.  You will need to download this [file from here](https://raw.githubusercontent.com/paulboal/hds5210-2023/main/week06/allergies.json) and then upload it into Google Colab to run these examples. It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE="allergies.json"
```

```python
import json
from pathlib import Path
HOME = str(Path.home())

```

This allergies.json file  contains allergy count data, this will be provide how many allergies factors contain in this file.

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
# Solution
def allergy_count(json_file):
    with open(json_file) as allergies :
        data = json.load(allergies)
        return len(data['entry'])
```

The above code provides total number of allergies file in alleries.json file.

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

This file have 4 numbers of allergies in 'entry'.

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
# Solution
def patient_count(ALLERGIES_FILE):
    with open (ALLERGIES_FILE) as af:

        data = json.load(af)

        all_patients = [ ]
        entry_array = data['entry']

        for entry in entry_array:
            all_patients.append(entry['resource'] ['patient'] ['display'])

        return len(set(all_patients))
```

This code provide how many patients are present in the alleries.json file, using of 'entry' have resource, patient and dispaly. (We are using the function patient_count to determine how many unique patients are included in this file.)

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

There are 2 numbers of patients in this json file.

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
# Solution
def allergy_per_patient(ALLERGIES_FILE):
    with open(ALLERGIES_FILE) as json_file:
        data = json.load(json_file)

    all_patients = [ ]

    entry_array = data['entry']

    for entry in entry_array:
        all_patients.append(entry['resource'] ['patient'] ['display'])

    allergy_per_patient_dict = dict()

    for patient in all_patients:
        if patient not in allergy_per_patient_dict:
            allergy_per_patient_dict[patient] = 1

        else:
            allergy_per_patient_dict[patient] += 1

    return allergy_per_patient_dict
```

This function provides the total number of allergies per patient as well as the number of patients who are effected by allergies.

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
# Solution
def allergy_list(ALLERGIED_FILE):
    with open (ALLERGIES_FILE)as json_file:
        data = json.load(json_file)

        all_patients_data = [ ]
        entry_array = data['entry']
    for entry in entry_array:
        patient_data = [ ]
        entry_resourse = entry['resource']
        patient_data.append(entry_resourse['patient'] ['display'])
        patient_data.append(entry_resourse['substance'] ['text'])
        patient_data.append(entry_resourse['reaction'][0]['manifestation'][0]['text'])
        all_patients_data.append(patient_data)
    return all_patients_data
```

This code will be provides total patient data including information and extractions names in json file. Allergies and reactions are included in entry_resource_substance aand display.

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
# Solution
def allergy_reaction(ALLERGIES_FILE, patient, substance):

    total_allergy_list = allergy_list(ALLERGIES_FILE)

    for allergy in total_allergy_list:
        if allergy[0] == patient and allergy[1] == substance:
            return allergy[2]
    return None
```

By reading a JSON file containing information about allergies, this tool determines whether a specific patient is allergic to a specific drug. If the patient has an allergy to the medicine, it will result in the appropriate response. Unless the patient has a medication or chemical allergy If it isn't listed in the allergy list, none

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

